![CC-BY-SA](https://mirrors.creativecommons.org/presskit/buttons/88x31/svg/by-sa.svg)


This notebook was created by [Bernardo Freitas Paulo da Costa](http://www.im.ufrj.br/bernardofpc),
and is licensed under Creative Commons BY-SA.

Antes de enviar este Teste, verifique que tudo está funcionando como esperado.
Por exemplo, **rode o código inteiro, do zero**.
Para isso, vá no menu, escolha _Kernel_, depois _Restart & Run All_.

Verifique, também, que você respondeu todas as questões:
* as questões de código têm `YOUR CODE HERE` (e você pode apagar o `raise NotImplemented` ao incluir sua resposta)
* as questões discursivas têm "YOUR ANSWER HERE".

---

# Teste 2: Mais bisseção

Vamos continuar estudando o método de bisseção, agora com garantias de erros absolutos e relativos ao mesmo tempo.

## Questão 1: uma função auxiliar

Quando a bisseção retorna uma "raiz", ela dá o ponto médio de um intervalo $[a,b]$, onde temos certeza que há uma raiz de $f$.

Escreva uma função `absrel(a,b)` que, dado um intervalo $[a,b]$, retorna o maior erro absoluto e o maior erro relativo que pode ocorrer ao considerar que a raiz é o ponto médio deste intervalo.

```python
import numpy as np
from numpy import infty
```

<!-- grade_id: absrel_fun, points: 0, grade: false, solution: true -->
```python
def absrel(a,b):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
absrel(10,11)
```

<!-- grade_id: absrel_test1, points: 2, grade: true, solution: false -->
```python
assert absrel(1,2) == (.5, .5)
assert absrel(10,11) == (0.5, 0.05)
```

<!-- grade_id: absrel_test2, points: 1, grade: true, solution: false -->
```python
assert absrel(-2,1) == (1.5, +infty)
```

Porquê o erro relativo do intervalo $[-2,1]$ é infinito?

<!-- grade_id: erro_relativo, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## Questão 2: Outros retornos na bisseção

Modifique o método da bisseção para terminar quando o erro absoluto **e** o erro relativo
forem menores do que (ou iguais às) tolerâncias dadas em argumento.

Retorne uma tripla contendo, nesta ordem:
- uma estimativa para a raiz $z$,
- uma estimativa para o erro absoluto $e_{abs}$,
- uma estimativa para o erro relativo $e_{rel}$.

<!-- grade_id: biss, points: 0, grade: false, solution: true -->
```python
def bissecao(f, a, b, abstol=1e-10, reltol=1e-10):
    """ Escreva uma documentação para explicar como este método funciona,
        qual o significado dos valores de entrada e dos de retorno. """
    # YOUR CODE HERE
    raise NotImplementedError()
```

Não se preocupe em obter valores iguais para as cotas superiores dos erros absoluto e relativo.

O importante é que estas estimativas sejam **coerentes** com os requisitos do método da bisseção:

- o erro absoluto real (ou seja, a diferença entre a raiz certa e a raiz calculada) deve ser menor do que a estimativa do erro absoluto,
- o erro relativo real deve ser menor do que a estimativa do erro relativo.

Note que é exatamente isso que vai ser testado em seguida!

```python
bissecao(np.cos, 1,5)
```

<!-- grade_id: biss_test1, points: 2, grade: true, solution: false -->
```python
x,err,rel = bissecao(np.sin,1,5)

# Testando que o valor retornado está perto da resposta, e satisfaz as próprias estimativas
assert abs(x - np.pi)/np.pi < rel
assert abs(x - np.pi) < err
```

<!-- grade_id: biss_test2, points: 1, grade: true, solution: false -->
```python
# Testando que sai pela razão certa
assert rel < 1e-10
assert err < 1e-10
```

```python
def f(x):
    return ((x**2 - 2)**2 - 2)**2 - 2 - x
```

<!-- grade_id: biss_test3, points: 2, grade: true, solution: false -->
```python
x,err,rel = bissecao(f,1.4,1.6,reltol=1e-12)

# Testando que o valor retornado está perto da resposta, e satisfaz as próprias estimativas
v = 1.532088886238
assert abs(x - v)/v < rel
assert abs(x - v) < err
```

<!-- grade_id: biss_test4, points: 1, grade: true, solution: false -->
```python
# Testando que sai pela razão certa
assert err < 1e-10
assert rel < 1e-12
```

## Questão 3: extremidades da bisseção

Modifique a sua implementação da bisseção para retornar uma lista com os intervalos considerados,
em ordem decrescente.

<!-- grade_id: biss_extr, points: 1, grade: true, solution: true -->
```python
def bissecao_extremidades(f,a,b, abstol=1e-10,reltol=1e-10):
    """ Escreva a documentação. """
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# Vendo alguns dos intervalos iniciais
l = bissecao_extremidades(np.sin,1,5)
for I in l[:5]:
    print(I)
```

<!-- grade_id: bisse_test1, points: 2, grade: true, solution: false -->
```python
lista_intervalos = bissecao_extremidades(np.sin,100,103)

Ifinal = lista_intervalos[-1]
(af, bf) = Ifinal

assert abs(bf - af) <= 2e-10
assert 35 <= len(lista_intervalos) <= 40
```

<!-- grade_id: bisse_test2, points: 2, grade: true, solution: false -->
```python
lista_intervalos = bissecao_extremidades(np.sin,100,103, abstol=1e-12, reltol=1e-12)

Ifinal = lista_intervalos[-1]
(af, bf) = Ifinal

assert abs(bf - af) <= 2e-12
assert abs(af - 32*np.pi) <= 2e-12
assert abs(bf - 32*np.pi) <= 2e-12
assert 40 <= len(lista_intervalos) <= 45
```

<!-- grade_id: bisse_test3, points: 2, grade: true, solution: false -->
```python
lista_intervalos = bissecao_extremidades(f,-0.5,0, abstol=1e-8, reltol=1e-12)

(af, bf) = lista_intervalos[-1]
x = (af+bf)/2

v = -0.4450418679126287
assert(abs(x - v)/v < 1e-12)
assert(abs(x - v) < 1e-8)
```

<!-- grade_id: bisse_test4, points: 1, grade: true, solution: false -->
```python
lista_intervalos = bissecao_extremidades(f,-0.5,0, abstol=1e-8, reltol=1e-12)

(af, bf) = lista_intervalos[-1]
x = (af+bf)/2

assert 40 <= len(lista_intervalos) <= 45
```

## Questão 4: Gráficos

Faça o gráfico do tamanho dos intervalos, em função do número de iterações,
para a bisseção do cosseno entre 0 e 3, com as tolerâncias padrão.

```python
import matplotlib.pyplot as plt
```

<!-- grade_id: graf1, points: 1, grade: true, solution: true -->
```python
a, b = 0, 3

# YOUR CODE HERE
raise NotImplementedError()

# Não apague a linha abaixo, ela faz os assert's dos gráficos darem certo
ax = plt.gca()
```

<!-- grade_id: graf_test1, points: 1, grade: true, solution: false -->
```python
assert len(ax.lines) == 1
assert len(ax.get_title()) >= 10
```

<!-- grade_id: graf_test2, points: 2, grade: true, solution: false -->
```python
linha = ax.lines[0]
assert 5e-11 < min(linha.get_ydata()) < 2e-10
assert max(linha.get_ydata()) == 3
```

Agora, faça um gráfico com
- o comprimento dos intervalos
- a distância entre a extremidade inferior e a raiz correta
- a distância entre a extremidade superior e a raiz correta

<!-- grade_id: graf2, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

ax = plt.gca()
```

<!-- grade_id: graf2_test1, points: 2, grade: true, solution: false -->
```python
assert len(ax.lines) == 3
assert len(ax.get_legend().texts) == 3
```

Repita o gráfico acima para o polinômio $f$ acima, com tolerância `1e-13`, tanto absoluta como relativa.

<!-- grade_id: graf3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

ax = plt.gca()
```

<!-- grade_id: graf3_test1, points: 1, grade: true, solution: false -->
```python
assert len(ax.lines) == 3
assert len(ax.get_legend().texts) == 3
```

<!-- grade_id: graf3_test2, points: 1, grade: true, solution: false -->
```python
for l in ax.lines:
    assert 40 <= len(l.get_xdata()) <= 45
```

O que você observa de comum nos gráficos com relação aos erros entre as extremidades dos intervalos?

<!-- grade_id: graf_analise, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE