![CC-BY-SA](https://mirrors.creativecommons.org/presskit/buttons/88x31/svg/by-sa.svg)


This notebook was created by [Bernardo Freitas Paulo da Costa](http://www.im.ufrj.br/bernardofpc),
and is licensed under Creative Commons BY-SA.

Antes de enviar este Teste, verifique que tudo está funcionando como esperado.
Por exemplo, **rode o código inteiro, do zero**.
Para isso, vá no menu, escolha _Kernel_, depois _Restart & Run All_.

Verifique, também, que você respondeu todas as questões:
* as questões de código têm `YOUR CODE HERE` (e você pode apagar o `raise NotImplemented` ao incluir sua resposta)
* as questões discursivas têm "YOUR ANSWER HERE".

---

# Teste 1: Bisseção

```python
import numpy as np
```

## Questão 1: Número de iterações da bisseção, e de chamadas à função `f`

Generalize o algoritmo da bisseção para retornar, nesta ordem:
- a (aproximação da) raiz,
- o número de bisseções (divisões por 2 do intervalo) feitas,
- o número de vezes que você chamou a função `f`.

Use como critério de parada a tolerância `tol`,
ou seja, o algoritmo termina quando:
- seja possível garantir que o erro (absoluto) da resposta ("em $x$") seja menor do que `tol`.

<!-- grade_id: cell-bfd2f7379cc308c4, points: 0, grade: false, solution: true -->
```python
def bissecao(f, a, b, tol=1e-8):
    z = (a + b) / 2
    niters = 1
    
    
    if (b - a) < tol:
        return z, niters
    
    
    if f(z)*f(a) < 0:
        tupla = bissecao(f, a, z, tol)
        
    else:
        tupla = bissecao(f, z, b, tol)
    
    raiz = tupla[0]
    niters += tupla[1]
    
    return raiz, niters, niters*2
```

### Testes simples

```python
def p2(x): return x**2 - 2
```

```python
bissecao(p2, 0, 2)
```

Output:
```
(1.414213564246893, 29, 58)
```

<!-- grade_id: cell-c62e828e23ba0c73, points: 1, grade: true, solution: false -->
```python
raiz, niters, ncalls = bissecao(p2, 0, 2)
assert abs(raiz - np.sqrt(2)) < 1e-8
assert 26 <= niters <= 30
```

<!-- grade_id: cell-3b33b89bcafe1954, points: 1, grade: true, solution: false -->
```python
raiz, niters, ncalls = bissecao(p2, 0, 2, tol=1e-3)
assert abs(raiz - np.sqrt(2)) < 1e-3
assert 10 <= niters <= 15
```

<!-- grade_id: cell-432c2ad26a7a7d7e, points: 1, grade: true, solution: false -->
```python
raiz, niters, ncalls = bissecao(p2, 0, 2, tol=1e-3)
assert abs(raiz - np.sqrt(2)) > 3e-4
```

<!-- grade_id: cell-721772ed44e72acb, points: 2, grade: true, solution: false -->
```python
# Testes escondidos aqui
```

### Testes mais legais

<!-- grade_id: cell-8d7ce0c5727c5838, points: 1, grade: true, solution: false -->
```python
pimeios, niters, ncalls = bissecao(np.cos, 0, 2)
assert abs(pimeios - np.pi/2) <= 1e-8
```

<!-- grade_id: cell-ff05771e68d86f63, points: 1, grade: true, solution: false -->
```python
menospimeios, niters, ncalls = bissecao(np.cos, -2, 0)
assert abs(menospimeios + np.pi/2) <= 1e-8
```

## Questão 2: várias iterações da bisseção

```python
print(f'{"raiz":^9} {"niters"} {"ncalls"}, {"erro":^11} {"tol":^10}')
for p in range(10,20):
    raiz, niters, ncalls = bissecao(p2, 0, 2, tol=2**(-p))
    print(f"{raiz:.7f} {niters:6d} {ncalls:6d}, {raiz-np.sqrt(2): .8f} {2**(-p):.8f}")
```

Output:
```
raiz    niters ncalls,    erro        tol    
1.4143066     13     26,  0.00009308 0.00097656
1.4141846     14     28, -0.00002899 0.00048828
1.4142456     15     30,  0.00003204 0.00024414
1.4142151     16     32,  0.00000153 0.00012207
1.4141998     17     34, -0.00001373 0.00006104
1.4142075     18     36, -0.00000610 0.00003052
1.4142113     19     38, -0.00000229 0.00001526
1.4142132     20     40, -0.00000038 0.00000763
1.4142141     21     42,  0.00000057 0.00000381
1.4142137     22     44,  0.00000010 0.00000191
```

O que você observa na tabela acima?  Explique.

<!-- grade_id: cell-2f026a9298f20d3b, points: 3, grade: true, solution: true -->
Um acontecimento interessante que pode ser evidenciado pela tabela, é que nem sempre um número de iterações menor apresentará mais erro, como é evidenciado pelo erro da iteração 16 (0.00000153) e o erro da iteração 17 (0.00001373). E isso ocorre algumas vezes na lista.

É claro que ,normalmente , quanto mais iterações menor o erro.

E Não me parece trivial o porque disso.

## Questão 3: Outros critérios de parada

Modifique a sua bisseção para ter 2 critérios de parada:
- quando o tamanho do intervalo ficar menor do que `tol`
- quando o valor da função ficar menor do que `err`

<!-- grade_id: cell-89ea2904b3c908eb, points: 0, grade: false, solution: true -->
```python
def bissecao(f, a, b, tol=1e-8, err=1e-8):
    z = (a + b) / 2
    niters = 1
    
    if abs(f(z)) <= err:
        return z, niters, niters*2
    
    if (b - a) <= tol:
        return z, niters, niters*2

    if f(z)*f(a) < 0:
        tupla = bissecao(f, a, z, tol)

    else:
        tupla = bissecao(f, z, b, tol)

    raiz = tupla[0]
    niters += tupla[1]
    
    return raiz, niters, niters*2
```

Pare e pense:

uma solução iterativa pode ficar mais fácil de adaptar para vários critérios de parada
se ela for escrita com `while True`.

### Testes

<!-- grade_id: cell-e5dbbdc7bb04bd36, points: 2, grade: true, solution: false -->
```python
raiz, niters, ncalls = bissecao(p2, 0, 2)
assert abs(p2(raiz) - 2) > 1e-8
assert abs(raiz - np.sqrt(2)) < 1e-8
```

<!-- grade_id: cell-65495ebe135fbf41, points: 1, grade: true, solution: false -->
```python
assert bissecao(np.cos, 0, 2) != bissecao(np.cos, 0, 2, err=0)
```

Output:
```
AssertionError                            Traceback (most recent call last)
Input In [17], in <cell line: 1>()
----> 1 assert bissecao(np.cos, 0, 2) != bissecao(np.cos, 0, 2, err=0)

AssertionError:
```

<!-- grade_id: cell-9446992c617f02eb, points: 1, grade: true, solution: false -->
```python
assert bissecao(np.cos, 0, 2) == bissecao(np.cos, 0, 2, tol=0)
```

O que os dois testes acima indicam sobre a bisseção para `np.cos`?

<!-- grade_id: cell-d8ce30660516f95d, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## Questão 4: Bisseção sem função

Escreva uma função que divide o intervalo ao meio, fica com o esquerdo,
e repete até o limite de precisão do computador.

Ela deve retornar o último intervalo não-trivial (como uma tupla), e o número de divisões.

<!-- grade_id: cell-dbe61010623d813f, points: 0, grade: false, solution: true -->
```python
def bisseção_esquerda(a,b):
    z = (a + b) / 2
    ndivs = 0
    
    if z == a:
        return (a, b), ndivs
    
    tupla = bisseção_esquerda(a, z)
    
    ndivs = tupla[1] + 1
    par = tupla[0]

    return par, ndivs
```

```python
(an, bn), ndivs = bisseção_esquerda(1,2)
assert 2e-16 <= bn - an <= 4e-16
assert ndivs == 52
print(an, bn)
```

Output:
```
1 1.0000000000000002
```

```python
(an, bn), ndivs = bisseção_esquerda(1,2)
assert 2e-16 <= bn - an <= 4e-16
assert ndivs == 52
```

<!-- grade_id: cell-365cb4391a9f65a4, points: 1, grade: true, solution: false -->
```python
(an, bn), ndivs = bisseção_esquerda(1,2)
assert 2e-16 <= bn - an <= 4e-16
assert ndivs == 52
```

<!-- grade_id: cell-179c89a09dfda258, points: 1, grade: true, solution: false -->
```python
(an, bn), ndivs = bisseção_esquerda(bn,2)
assert 2e-16 <= bn - an <= 4e-16
assert ndivs == 51
```

Use esta função para calcular o próximo número de ponto flutuante.

<!-- grade_id: cell-ace3480d37572adb, points: 0, grade: false, solution: true -->
```python
def próximo_float(x):
    prox = x + 1
    tupla = bisseção_esquerda(x, prox)
    return tupla[0][1]
```

```python
próximo_float(1)
```

Output:
```
1.0000000000000002
```

<!-- grade_id: cell-8f75b70a7dbab81a, points: 1, grade: true, solution: false -->
```python
nextone = próximo_float(1)
assert nextone > 1
assert (1+nextone)/2 == 1
```

<!-- grade_id: cell-b39b0becc0922ad6, points: 1, grade: true, solution: false -->
```python
nextzero = próximo_float(0)
assert nextzero > 0
assert nextzero/2 == 0
```

<!-- grade_id: cell-c902bc7ab832c0d0, points: 1, grade: true, solution: false -->
```python
nextnext = próximo_float(próximo_float(1))
assert nextnext > 1
assert 1 < (nextnext+1)/2 < nextnext
```

<!-- grade_id: cell-96bd65cd68e38dc0, points: 1, grade: true, solution: false -->
```python
maisum = próximo_float(-1.435)
assert maisum > -1.435
assert (maisum - 1.435)/2 == -1.435
```

## Questão 5: bisseção alternante

Escreva uma função que escolhe o lado esquerdo, e depois o lado direito, e repete, até acabar.

<!-- grade_id: cell-3a6605b0c2fff49d, points: 0, grade: false, solution: true -->
```python
def bisseção_alt(a,b):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-d1906fd95ebb85a3, points: 1, grade: true, solution: false -->
```python
(a,b), ndivs = bisseção_alt(1,2)
assert 1 < a < b < 2
assert ndivs == 52
```

<!-- grade_id: cell-1433725a5fa659f7, points: 1, grade: true, solution: false -->
```python
(a,b), ndivs = bisseção_alt(1,4)
assert 1 < a < b < 4
assert ndivs == 53
assert (a+b)/2 == b
```

```python
for (a0,b0) in [(1,2), (1,4), (2,4), (1/3,3/5)]:
    (a,b), ndivs = bisseção_alt(a0,b0)
    print(f"{a0:.5f} {b0:.5f} ---> {a:.17f} {b:.17f}")
```

O que esta função parece calcular?  Explique.

<!-- grade_id: cell-4d30f5821b51b3d7, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: cell-36e18b27847e0f36, points: 1, grade: true, solution: false -->
```python
(a,b), ndivs = bisseção_alt(-1,2)
assert -1 < a < b < 2
assert (a+b)/2 == b
assert ndivs > 1000
```

Esta iteração levou mais de 1000 bisseções.  Explique o que aconteceu.

<!-- grade_id: cell-203e11f457638867, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE