![CC-BY-SA](https://mirrors.creativecommons.org/presskit/buttons/88x31/svg/by-sa.svg)


This notebook was created by [Bernardo Freitas Paulo da Costa](http://www.im.ufrj.br/bernardofpc),
and is licensed under Creative Commons BY-SA.

Antes de enviar este Teste, verifique que tudo está funcionando como esperado.
Por exemplo, **rode o código inteiro, do zero**.
Para isso, vá no menu, escolha _Kernel_, depois _Restart & Run All_.

Verifique, também, que você respondeu todas as questões:
* as questões de código têm `YOUR CODE HERE` (e você pode apagar o `raise NotImplemented` ao incluir sua resposta)
* as questões discursivas têm "YOUR ANSWER HERE".

---

```python
import numpy as np
import matplotlib.pyplot as plt
```

# Regressão periódica

## Dados

Nesta primeira parte iremos ajustar uma curva aos dados de demanda elétrica fornecidos pelo ONS,
para o período entre os dias *16.04.2018 - 16.05.2018*.

```python
cch = np.load('Demanda_16.04.2018-16.05.2018.npz')
ys = cch['Demanda']
ts = np.arange(len(ys))

plt.figure(figsize=(15,5))
plt.plot(ts,ys)
plt.title("Curva de carga horária")
plt.ylabel("MWh")
plt.xlabel("Intervalos de 1h a partir de 2018-04-16 00:00");
```

Vamos pegar um intervalo com dias úteis.

```python
tmax = 120

plt.plot(ts[:tmax],ys[:tmax])
plt.title("Curva de carga horária")
plt.ylabel("MWh")
plt.xlabel("Intervalos de 1h a partir de 2018-04-16 00:00");
```

# Parte 1: Ajuste simples

## Questão 1: Matriz da regressão

O que salta aos olhos é o comportamento aproximadamente periódico,
logo uma primeira idéia é ajustar um modelo com funções também periódicas.

Um ajuste simples é o seguinte:
$$y(t) = c_0 + c_1\sin\left(\frac{2\pi t}{24}\right) + c_2\cos\left(\frac{2\pi t}{24}\right)$$

Explique porque o denominador 24 faz sentido para estes dados que temos.

<!-- grade_id: cell-1ff160f6f6414643, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

Escreva a função `make_trig_matrix(ts,p)`, que monta a matriz de regressão $X$ tal que $Xc = y$ para o modelo proposto acima para `ts` e com período `p`.

**Observações:**
- A função constante está na base;
- A ordem das colunas importa para os asserts (mas não para o ajuste, se você sabe o que faz).
  Assim, as colunas devem vir na ordem das funções: constante, seno, cosseno.

<!-- grade_id: cell-a46db828fddba423, points: 0, grade: false, solution: true -->
```python
def make_trig_matrix(ts,p):
    """ Retorna a matriz de regressão para o modelo y(t) = c0 + c1*sin(2*pi*t/p) + c2*cos(2*pi*t/p). """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-9f2d672d19f0269e, points: 1, grade: true, solution: false -->
```python
tss = np.linspace(-2*np.pi,2*np.pi,51)
X = make_trig_matrix(tss,np.pi)

assert all([np.isclose(X[i][0],1) for i in range(10)])
```

<!-- grade_id: cell-f609d0ed36084424, points: 2, grade: true, solution: false -->
```python
# Testes de periodicidade

assert np.allclose(X[0][1],X[-1][1])
assert np.allclose(X[0][1],X[len(tss)//2][1])
assert np.allclose(X[0][2],X[-1][2])
assert np.allclose(X[0][2],X[len(tss)//2][2])
```

## Questão 2: Gráficos

Faça o gráfico tanto dos dados como do ajuste do modelo,
usando apenas as 120 primeiras observações, que correspondem aos primeiros cinco dias úteis.

<!-- grade_id: cell-ac3a21418fecbfda, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

ax = plt.gca()
plt.show()
```

<!-- grade_id: cell-009fef7da3f75e1f, points: 2, grade: true, solution: false -->
```python
assert ax.title.get_text() != ""
assert len(ax.lines) == 2
assert len(ax.legend().texts) == 2

yss = [ax.lines[i].get_ydata() for i in range(2)]

assert 73000 <= np.max(yss) <= 73800
assert 49000 <= np.min(yss) <= 50000
```

```python
ax = None
```

Plote também o residual (erro de ajuste)!

<!-- grade_id: cell-2d246d7cd0dffdd7, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

ax = plt.gca()
plt.show()
```

<!-- grade_id: cell-70943487ee6edb3c, points: 2, grade: true, solution: false -->
```python
assert ax.title.get_text() != ""
assert len(ax.lines) == 1

```

```python
ax = None
```

Esse gráfico reflete alguma tendência nos dados que difere do ajuste?
Se sim qual função poderia ser acrescentada a base do modelo para tentar capturar essa tendência?

<!-- grade_id: cell-07fc4ad29c509557, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

O que mais você percebe no gráfico de erros?

<!-- grade_id: cell-45ed8fa6dad52463, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

## Questão 3: Uma base maior

Agora, vamos tentar um ajuste com mais funções periódicas (também ditas _harmônicas_):
$$y(t) = \alpha_0 +  \sum_{i=1}^k\alpha_i \sin\left(\frac{2\pi it}{24}\right) + \sum_{i=1}^k\beta_i \cos\left(\frac{2\pi it}{24}\right).$$

Para te auxiliar nessa tarefa escreva a função
- `make_full_trig_matrix(ts,p,k)`: Semelhante a `make_trig_matrix`, mas retorna a matriz para o ajuste acima,
  correspondente ao conjunto de harmônicos de $i=1$ até $i = k$.

<!-- grade_id: cell-126e0ede6c80509a, points: 0, grade: false, solution: true -->
```python
def make_full_trig_matrix(ts,p,k):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-2f5c864ca5ccb8c5, points: 2, grade: true, solution: false -->
```python
# Teste de sanidade
tss = np.arange(0,100)

X = make_full_trig_matrix(tss, 24, 1)
Y = make_trig_matrix(tss, 24)

for i in range(5):
    for j in range(3):
        assert np.isclose(X[i][j],Y[i][j])
```

<!-- grade_id: cell-afc68507ca22634a, points: 1, grade: true, solution: false -->
```python
# Teste de dimensão
tss = np.arange(0,100)

X = make_full_trig_matrix(tss, 24, 3)
assert np.shape(X) == (100,7)
```

## Questão 4: Gráficos de novo

Faça o gráfico do ajuste para $k = 2$, novamente para os 5 primeiros dias úteis.

<!-- grade_id: cell-3880482d179915d1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

ax = plt.gca()
plt.show()
```

<!-- grade_id: cell-a06ed347faa66435, points: 2, grade: true, solution: false -->
```python
assert ax.title.get_text() != ""
assert len(ax.lines) == 2
assert len(ax.legend().texts) == 2

for i in range(2):
    yss = ax.lines[i].get_ydata()
    assert 71000 <= np.max(yss) <= 73300
    assert 49000 <= np.min(yss) <= 51300
```

```python
ax = None
```

E não se esqueça do residual.

<!-- grade_id: cell-8e3f9cb66f1ca4ab, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

ax = plt.gca()
plt.show()
```

<!-- grade_id: cell-7955ff98c82afd98, points: 1, grade: true, solution: false -->
```python
assert ax.title.get_text() != ""
assert len(ax.lines) == 1

```

```python
ax = None
```

Agora, faça o gráfico do erro para $k = 5$, $6$ e $10$, lado a lado.

**Dica:** Faça uma função `graph_err(k, tmax)` que faz o gráfico do erro do ajuste em função de `k` e `tmax`,
e use `subplot(1,3,j)` para selecionar o $j$-ésimo eixo.
(como é inspirada do MATLAB, essa função conta a partir de 1, não zero como padrão em python).

<!-- grade_id: cell-1129a6a77b1ea8ce, points: 2, grade: true, solution: true -->
```python
fig, axs = plt.subplots(ncols=3, figsize=(18,4))

# YOUR CODE HERE
raise NotImplementedError()

plt.show()
```

O que você muda e o que não muda ao aumentar $k$?

<!-- grade_id: cell-0e2470c89a85b56d, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

# Parte 2: Ajuste misto

## Questão 5: Polinômios e funções trigonométricas

Vamos adicionar polinômios à base, e comparar o erro de regressão obtido com uma base um pouco maior.

Comece implementando a função `make_poly_matrix(ts, d_max)`,
que retorna a matriz de ajuste polinomial de grau até $d_\max$.

Para não duplicar a coluna de 1's na matriz do modelo completo,
a função `make_poly_matrix` não retorna a coluna correspondente às funções constantes,
mas apenas aos polinômios $X$, $X^2$, ..., $X^{d_\max}$.

<!-- grade_id: cell-ec8fe8f07c2424db, points: 0, grade: false, solution: true -->
```python
def make_poly_matrix(ts, dmax):
    assert dmax >= 1
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-bd996f7715ed7887, points: 2, grade: true, solution: false -->
```python
tss = np.arange(1,5)

X = make_poly_matrix(tss,3)

assert np.allclose(X[0],[1,1,1])
assert np.allclose(X[1],[2,4,8])
assert np.allclose(X[2],[3,9,27])
assert np.allclose(X[3],[4,16,64])
```

Agora, faça uma função que retorna a matriz do ajuste misto

<!-- grade_id: cell-354b03c14e91f0dc, points: 0, grade: false, solution: true -->
```python
def make_trig_poly(ts, p, k, dmax):
    if dmax == 0:
        return make_full_trig_matrix(ts, p, k)
    else:
        # YOUR CODE HERE
        raise NotImplementedError()
```

## Questão 6: Mais gráficos ainda

Trace, num mesmo eixo, os erros correspondentes aos seguintes ajustes:
- Trigonométrico original, com apenas 1 harmônico
- Trigonométrico com 6 harmônicos
- Misto com 1 harmônico e um polinômio de grau 1
- Misto com 6 harmônicos e polinômio de grau 1
- Misto com 1 harmônico e polinômios de grau até 5
- Misto com 6 harmônicos e polinômios de grau até 5

<!-- grade_id: cell-9555d040a56e5596, points: 0, grade: false, solution: true -->
```python
def graph_err(ts, ys, p, k, dmax):
    """Função auxiliar para gráfico de erros"""
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
plt.figure(figsize=(15,5))

subts = ts[:120]
subys = ys[:120]

graph_err(subts, subys, 24, 1, 0)
graph_err(subts, subys, 24, 6, 0)
graph_err(subts, subys, 24, 1, 1)
graph_err(subts, subys, 24, 6, 1)
graph_err(subts, subys, 24, 1, 5)
graph_err(subts, subys, 24, 6, 5)

plt.axhline(color='black', linestyle='--', linewidth=1)
ax = plt.gca()
plt.show()
```

<!-- grade_id: cell-053dd1cf5f0db60b, points: 2, grade: true, solution: false -->
```python
assert ax.title.get_text() != ""
assert len(ax.lines) == 7
assert len(ax.legend().texts) == 6
```

O que você observou? Que melhora é obtida com o acréscimo dos polinômios da base?

<!-- grade_id: cell-5f8f1cd0f637fda7, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

Faça agora dois gráficos de erro, para separar em dois grupos.

<!-- grade_id: cell-afea4cc93bf320a2, points: 2, grade: true, solution: true -->
```python
fig, axs = plt.subplots(ncols=2, figsize=(13,4), sharey=True)

subts = ts[:120]
subys = ys[:120]

# YOUR CODE HERE
raise NotImplementedError()
```

- Quais são os grupos que você escolheu?
- Porquê?

<!-- grade_id: cell-b7f608fd47648c78, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE