![CC-BY-SA](https://mirrors.creativecommons.org/presskit/buttons/88x31/svg/by-sa.svg)


This notebook was created by [Bernardo Freitas Paulo da Costa](http://www.im.ufrj.br/bernardofpc),
and is licensed under Creative Commons BY-SA.

Antes de enviar este Teste, verifique que tudo está funcionando como esperado.
Por exemplo, **rode o código inteiro, do zero**.
Para isso, vá no menu, escolha _Kernel_, depois _Restart & Run All_.

Verifique, também, que você respondeu todas as questões:
* as questões de código têm `YOUR CODE HERE` (e você pode apagar o `raise NotImplemented` ao incluir sua resposta)
* as questões discursivas têm "YOUR ANSWER HERE".

---

```python
import numpy as np
import matplotlib.pyplot as plt
from time import time
```

# Resolvendo sistemas especiais

## Questão 1: Algoritmo

Uma matriz é dita tridiagonal quando só possui entradas diferentes de zero na diagonal,
e imediatamente acima ou abaixo dela.

Explique porque, ao fazer a eliminação, o algoritmo só precisa operar nas três "diagonais centrais".

<!-- grade_id: cell-671f1d7ad1aebc5e, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

Dê um exemplo de uma matriz $4 \times 4$, com muitos zeros, mas para a qual o algoritmo de eliminação
vai ter que acessar e calcular nas entradas correspondentes a estes zeros.

<!-- grade_id: cell-d5ca797d17ccc3d9, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

Agora, escreva a função `tridiag_solve` para resolver sistemas tridiagonais, ou seja,
aqueles cuja matriz de coeficientes é tridiagonal.

Esta função deve realizar o algoritmo de eliminação,
mas, para ter uma boa performance, opere apenas nas entradas das diagonais centrais.

<!-- grade_id: cell-462b708f5e397ecf, points: 0, grade: false, solution: true -->
```python
def tridiag_solve(A, b):
    A = np.copy(np.asarray(A, dtype=np.float64))
    b = np.copy(np.asarray(b, dtype=np.float64))
    n,m = np.shape(A)
    assert m == len(b)
    assert n == m, "Matrix must be square"
    # YOUR CODE HERE
    raise NotImplementedError()
```

Qual o papel do primeiro assert?

<!-- grade_id: cell-7a4214ebd68e358c, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: cell-65f2e07cb2fa5887, points: 1, grade: true, solution: false -->
```python
A = [[1, 2, 0], [2, 1, -1], [0, 1, -2]]
b = [1,1,1]

x1 = np.linalg.solve(A,b)
x2 = tridiag_solve(A, b)
assert np.allclose(x1, x2, atol=1e-14, rtol=1e-14)
```

<!-- grade_id: cell-20bb3249b0db5e25, points: 1, grade: true, solution: false -->
```python
np.random.seed(9)
A = np.diag(range(3,8)) + np.diag([-1,-1,-2,-2], k=1) + np.diag([-2,-2,-1,-1], k=-1)
b = np.random.randn(5)

x1 = np.linalg.solve(A, b)
x2 = tridiag_solve(A, b)
assert np.allclose(x1, x2, atol=1e-13, rtol=1e-13)
```

Se você não tivesse acesso a `np.linalg.solve`, como você poderia testar sua função?
Explique abaixo, e dê um exemplo de teste na caixa seguinte.

<!-- grade_id: cell-880f3f95c45d504d, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: cell-e359c03c250b063f, points: 1, grade: true, solution: true -->
```python
# Exemplo de teste
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-077e432dff6201e1, points: 1, grade: true, solution: false -->
```python
np.random.seed(9)
A = np.random.randn(5,5)
b = [1,2,3,4,5]


x1 = np.linalg.solve(A, b)
x2 = tridiag_solve(A, b)
assert not np.allclose(x1, x2, atol=1e-3, rtol=1e-3)
```

Explique o resultado dos testes acima.

<!-- grade_id: cell-2ee20344f5d12125, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

## Questão 2: Tempos

Para cada um dos `ns` dados a seguir, monte um sistema tridiagonal $n \times n$,
e calcule o tempo (usando `time`) para:
1. Criar a matriz A e o vetor b (podem ser aleatórios, podem não ser)
2. Resolver o sistema usando `tridiag_solve`

Armazene os tempos em duas listas, `t_create` e `t_solve`.

Obs: a caixa abaixo pode usar algo como 500M de RAM, por conta das matrizes 4k
(ou até mais, dependendo de **como** você criar as matrizes).
Assim, para testar, diminua a amplitude do logspace, terminando em $2^{10}$ em vez de $2^{12}$.
Se você não conseguir fazer funcionar até $2^{12}$,
reduza o maior valor do `logspace`, mas não reduza demais.

Obs2: A caixa abaixo deve ser executada em menos de 60 segundos, para não dar `TimeOut`.

<!-- grade_id: cell-84a663322d8e87f5, points: 2, grade: true, solution: true -->
```python
%%time

ns = np.logspace(4,12, base=2, num=28, dtype=int)
t_create = []
t_solve  = []
for i, n in enumerate(ns):
    # YOUR CODE HERE
    raise NotImplementedError()
```

Agora, faça um gráfico destes tempos, em função de $n$.

<!-- grade_id: cell-a67d25ce0784d87f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

ax = plt.gca()
```

<!-- grade_id: cell-e3e0f254799598f9, points: 1, grade: true, solution: false -->
```python
assert len(ax.lines) == 2
assert len(ax.legend().texts) == 2
assert ax.get_title() != ""
```

<!-- grade_id: cell-b0b81c2b16c50f3a, points: 1, grade: true, solution: false -->
```python
for l in ax.lines:
    xs = l.get_xdata()
    assert min(xs) == 2**4
    assert 2**9 <= max(xs) <= 2**13
```

```python
ax = None
```

Qual etapa demora mais?
Porquê você acha que isso acontece?

<!-- grade_id: cell-7d10cf0ddf503610, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

Agora, faça novos gráficos, para tentar descobrir a taxa de crescimento destes tempos:
se é linear, quadrático, exponencial, ... em função de $n$.

E se são a mesma taxa ou não!

<!-- grade_id: cell-0ac7be496b4ab4a4, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Comente

<!-- grade_id: cell-c5b389d9a6bae124, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## Questão 3: Sistemas pentadiagonais

Agora, suponha que a matriz $A$ é pentadiagonal, ou seja, as entradas não-nulas
podem estar até 2 linhas acima ou abaixo da diagonal principal.
Generalize o seu programa para resolver sistemas pentadiagonais.

<!-- grade_id: cell-41937ff612288545, points: 0, grade: false, solution: true -->
```python
def pentadiag_solve(A, b):
    A = np.copy(np.asarray(A, dtype=np.float64))
    b = np.copy(np.asarray(b, dtype=np.float64))
    n,m = np.shape(A)
    assert m == len(b)
    assert n == m, "Matrix must be square"
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-2e48084bf174185d, points: 1, grade: true, solution: false -->
```python
np.random.seed(9)
A = np.diag(range(3,8)) + np.diag([-1,-1,-2,-2], k=1) + np.diag([-2,-2,-1,-1], k=-1)
b = np.random.randn(5)

x2 = tridiag_solve(A, b)
x3 = pentadiag_solve(A, b)
assert np.allclose(x3, x2, atol=1e-13, rtol=1e-13)
```

<!-- grade_id: cell-ea99873bf680925a, points: 2, grade: true, solution: false -->
```python
np.random.seed(10)
A = np.diag(range(3,9)) + np.diag([-1,-1,-2,-2], k=2) + np.diag([-2,-2,-1,-1], k=-2)
b = np.random.randn(6)

x1 = np.linalg.solve(A, b)
x3 = pentadiag_solve(A, b)
assert np.allclose(x1, x3, atol=1e-13, rtol=1e-13)
```

<!-- grade_id: cell-a9b89660ea1f5a11, points: 1, grade: true, solution: false -->
```python
np.random.seed(10)
A = np.diag(range(3,9)) + np.diag([-1,-1,0,-2,-2], k=1) + np.diag([-2,-2,-1,-1], k=-2)
b = np.random.randn(6)

x1 = np.linalg.solve(A, b)
x3 = pentadiag_solve(A, b)
assert np.allclose(x1, x3, atol=1e-13, rtol=1e-13)
```

Repita os cálculos de tempo, para resolver sistemas de dimensão até $2^{12}$.

<!-- grade_id: cell-2e5b4eed7bb22b4b, points: 1, grade: true, solution: true -->
```python
ns = np.logspace(4,12, base=2, num=28, dtype=int)
t_create = []
t_solve  = []
for i, n in enumerate(ns):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-aa08140310aa9c71, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

ax = plt.gca()
```

<!-- grade_id: cell-daef85db8ad2ed47, points: 1, grade: true, solution: false -->
```python
assert len(ax.lines) == 2
assert len(ax.legend().texts) == 2
assert ax.get_title() != ""
```

O que mudou?

<!-- grade_id: cell-7e4a3bc07ef66e1e, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

## Questão 4: Velocidade

Seu método tridiagonal é mais rápido do que `np.linalg.solve` para matrizes $10 \times 10$?
Para qual $n$, aproximadamente, seu método fica mais rápido?

<!-- grade_id: cell-cc0d30be0b997cb8, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

E como é a qualidade da solução?
Faça cálculos na caixa abaixo, e comente em seguida.

<!-- grade_id: cell-b1b8d36719f10882, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-4f62e86e0f71ebab, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE