![CC-BY-SA](https://mirrors.creativecommons.org/presskit/buttons/88x31/svg/by-sa.svg)


This notebook was created by [Bernardo Freitas Paulo da Costa](http://www.im.ufrj.br/bernardofpc),
and is licensed under Creative Commons BY-SA.

Antes de enviar este Teste, verifique que tudo está funcionando como esperado.
Por exemplo, **rode o código inteiro, do zero**.
Para isso, vá no menu, escolha _Kernel_, depois _Restart & Run All_.

Verifique, também, que você respondeu todas as questões:
* as questões de código têm `YOUR CODE HERE` (e você pode apagar o `raise NotImplemented` ao incluir sua resposta)
* as questões discursivas têm "YOUR ANSWER HERE".

---

```python
import numpy as np
import matplotlib.pyplot as plt
```

# Parte 1: Projetores e Refletores

## Questão 1: Projeção ortogonal

Escreva uma função `proj(v)` que retorna **uma função**
que calcula a projeção ortogonal de um vetor qualquer $u$  em $v$.

<!-- grade_id: cell-923bb846ded32913, points: 0, grade: false, solution: true -->
```python
def proj(v):
    vv = np.asarray(v)
    def aux(u):
        u = np.asarray(u)
        numerador = u.T@vv
        denominador = vv.T@vv
        
        coeficiente = numerador/denominador
        return coeficiente*vv
    return aux
```

<!-- grade_id: cell-7c5f82e15593a689, points: 1, grade: true, solution: false -->
```python
f2 = proj([0,1])
assert np.allclose(f2([0,2]), [0,2])
assert np.allclose(f2([1,0]), [0,0])
```

<!-- grade_id: cell-b4495ff1baf43936, points: 2, grade: true, solution: false -->
```python
f3 = proj([0,1,2])
assert np.allclose(f3([0,1,1]), [0, 0.6, 1.2])
assert np.allclose(f3([2,1,0]), [0, 0.2, 0.4])
```

<!-- grade_id: cell-c5454b4c058650b1, points: 1, grade: true, solution: false -->
```python
fbig = proj([1,1,-2,3,-1,2])
assert np.allclose(4*fbig([1,1,1,1,0,1]), [1,1,-2,3,-1,2])
```

## Questão 2: Matriz da projeção

Escreva uma função `matr_proj(v)` que calcula a matriz de projeção no vetor $v$.

<!-- grade_id: cell-9392babacf493311, points: 0, grade: false, solution: true -->
```python
def matr_proj(v):
    # YOUR CODE HERE
    n = len(v)
    matriz = []
    projv = proj(v)

    for i in range(n):
        ei = np.zeros(n)    # Crio a base canônica
        ei[i] = 1
        
        ei_new = projv(ei)   # Aplico na base canônica
        matriz.append(ei_new)
    
    return np.array(matriz)
```

<!-- grade_id: cell-ee8eb79b77e28971, points: 2, grade: true, solution: false -->
```python
M2 = matr_proj([1,2])
assert np.allclose(M2, [[0.2, 0.4], [0.4, 0.8]])
```

<!-- grade_id: cell-daaf9d49d775c453, points: 2, grade: true, solution: false -->
```python
M5 = matr_proj([1,2,0,-3,1])
assert np.shape(M5) == (5,5)
assert abs(M5[3,3] - 0.6) < 1e-15
```

<!-- grade_id: cell-ca35206b552682dd, points: 1, grade: true, solution: false -->
```python
v = [1,-2]
u = [1,1]

M = matr_proj(v)
f = proj(v)

assert np.allclose(M@u, f(u))
```

Explique o assert acima

<!-- grade_id: cell-7a7c24d916217af9, points: 2, grade: true, solution: true -->
Ele se certifica de que, usar o produto escalar pra fazer a projeção e multiplicar a matriz projeção dão o mesmo resultado.

## Questão 3: Tempos

Compare o tempo que a função de projeção leva com o tempo que a multiplicação pela matriz de projeção leva,
em vetores de tamanho 80.

Dica: use `%timeit`, o tempo deve estar na ordem de **micro**segundos

<!-- grade_id: cell-aceeee77766c6b8b, points: 2, grade: true, solution: true -->
```python
v = np.random.rand(80)
u = np.random.rand(80)
# YOUR CODE HERE
%timeit matr_proj(v)@u
%timeit proj(v)(u)
```

Output:
```
449 µs ± 7.83 µs per loop (mean ± std. dev. of 7 runs, 1,000 loops each)
4.58 µs ± 105 ns per loop (mean ± std. dev. of 7 runs, 100,000 loops each)
```

Faça o mesmo, em vetores de tamanho 400.

<!-- grade_id: cell-fd7fd642aae03fdd, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
v = np.random.rand(400)
u = np.random.rand(400)
# YOUR CODE HERE
%timeit matr_proj(v)@u
%timeit proj(v)(u)
```

Output:
```
3.33 ms ± 256 µs per loop (mean ± std. dev. of 7 runs, 100 loops each)
5.06 µs ± 168 ns per loop (mean ± std. dev. of 7 runs, 100,000 loops each)
```

E de tamanho 2000...

<!-- grade_id: cell-545e5d7f4be7fef9, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
v = np.random.rand(2000)
u = np.random.rand(2000)
# YOUR CODE HERE
%timeit matr_proj(v)@u
%timeit proj(v)(u)
```

Output:
```
48 ms ± 2.53 ms per loop (mean ± std. dev. of 7 runs, 10 loops each)
6.58 µs ± 229 ns per loop (mean ± std. dev. of 7 runs, 100,000 loops each)
```

Explique porque isso acontece.

<!-- grade_id: cell-bbd7408d759f0f83, points: 2, grade: true, solution: true -->
Isso se deve a como construí a função que cria a matriz. Ele cria a base canônica, aplica a transformação nela e dessa informação extrai a matriz, pra depois aplicá-la ao vetor.
A projeção já trabalha apenas com a norma dos vetores e há a multiplicação pelo vetor.

## Questão 4: Projetores parciais

Suponha agora que temos um vetor $u \in R^{n+m}$,
e vamos projetar apenas as $m$ últimas coordenadas na direção do vetor $v \in R^m$
(as $n$ primeiras coordenadas permanecem inalteradas).
Generalize a sua função `proj(v)` para que a função retornada funcione para qualquer vetor $u$ de dimensão maior ou igual a $m$.

<!-- grade_id: cell-c2f353e01fe623c4, points: 0, grade: false, solution: true -->
```python
def proj_gen(v):
    m = len(v)
    def aux(u):
        assert len(u) >= m
        # YOUR CODE HERE
        uu = u[-m:] # Pegará as m últimas coordenadas

        projecao = proj(v)  # Função que projeta em v

        uu_new_m = projecao(uu) # Vetor uu projetado em v
        
        u[-m:] = uu_new_m   # Muda as M coordenadas que foram projetadas

        return u

    return aux
```

<!-- grade_id: cell-6cf4a0bdec853f70, points: 1, grade: true, solution: false -->
```python
g2 = proj_gen([1,1])
assert np.all(g2([-1,2.]) == [1/2, 1/2])
```

<!-- grade_id: cell-18d142fd066c10b6, points: 1, grade: true, solution: false -->
```python
g2 = proj_gen([1,1])
assert np.all(g2([6,5,4,3,2,1.]) == [6. , 5. , 4. , 3. , 1.5, 1.5])
```

<!-- grade_id: cell-7833e56994c906fb, points: 1, grade: true, solution: false -->
```python
g2 = proj_gen([1,-3])
assert np.allclose(g2([6,5,4,3,2,1.]), [ 6. ,  5. ,  4. ,  3. , -0.1,  0.3])
```

Que outros testes você faria para verificar que a sua função `proj_gen` está funcionando como imaginado?

Dê, na caixa abaixo, a justificativa para os testes,
e nas duas caixas seguintes, dois exemplos de cada um.

<!-- grade_id: cell-0177712a1b2ae519, points: 2, grade: true, solution: true -->
Poderíamos verificar se essa projeção está batendo com as outras formas de projeção que fizemos.

<!-- grade_id: cell-67502a09a665030e, points: 1, grade: true, solution: true -->
```python
# Testes 1
v = np.random.rand(80)
u = np.random.rand(80)

f = proj(v)
f2 = proj_gen(v)

assert np.allclose(f2(u), f(u))

v2 = np.random.rand(100)
u2 = np.random.rand(100)

f = proj(v2)
f2 = proj_gen(v2)

assert np.allclose(f2(u2), f(u2))
```

<!-- grade_id: cell-8b74a28008d4ae2b, points: 1, grade: true, solution: true -->
```python
# Testes 2
v = np.random.rand(80)
u = np.random.rand(80)

M = matr_proj(v)
f = proj_gen(v)

assert np.allclose(M@u, f(u))

v2 = np.random.rand(100)
u2 = np.random.rand(100)

M = matr_proj(v2)
f = proj_gen(v2)

assert np.allclose(M@u2, f(u2))

```

## Questão 5: Reflexões

Podemos construir uma reflexão $R_v$ a partir de uma projeção $P_v$ pela seguinte fórmula:
$$
  R_v(u) = u - 2(u - P_v(u)).
$$

Explique o significado do termo $u - P_v(u)$ na fórmula acima.

<!-- grade_id: cell-1f959cf7709080af, points: 2, grade: true, solution: true -->
Para facilitar pensaremos no caso $R^3$, onde queremos refletir uma reta com relação a um plano.

Temos a mesma distância entre: $u$ e $P_v(u)$, e $R_v(u)$ e $P_v(u)$.

E nesse caso é fácil visualizar que se pegarmos essa distância e diminuímos ela duas vezes do nosso vetor original teremos a nossa reflexão.

Simplifique a fórmula (claro!) e implemente a função `refl(v)` que faz a reflexão "segundo o vetor $v$".

<!-- grade_id: cell-70c58bc6a6c6cd3c, points: 0, grade: false, solution: true -->
```python
def refl(v):
    vv = np.asarray(v)
    def aux(u):
        u = np.array(u)

        projecao = np.array(proj(vv)(u))
        reflexao = 2*projecao - u
        
        return reflexao
    
    return aux
```

```python
r2 = refl([1,1])
print(r2([3, 1]))
```

Output:
```
[1. 3.]
```

<!-- grade_id: cell-5cc3fd8dddd6d9bb, points: 1, grade: true, solution: false -->
```python
r2 = refl([1,1])
assert np.allclose(r2([3,1]), [1,3])
```

<!-- grade_id: cell-7813c2f1678ec5af, points: 1, grade: true, solution: false -->
```python
r4 = refl([1,3,2,1])
assert np.allclose(r4([3,3,3,3]), [-0.2, 5.4, 2.6, -0.2])
```

## Questão 6: Refletindo a reflexão

Uma reflexão é uma **involução**: aplicada duas vezes, ela volta ao início.

Mostre (algebricamente) que a fórmula que demos para a reflexão de fato produz uma involução.
Ou seja, mostre que, dados $v$ e $u$ quaisquer, $R_v(R_v(u)) = u$.

<!-- grade_id: cell-39866ae20579c657, points: 2, grade: true, solution: true -->
$R_v(u) = 2P_v(u) - u $ --> $ u = 2P_v(u) - R_v(u)$

Jogando o $R_v$ dentro dele mesmo:

$R_v(R_v(u)) = 2P_v(R_v(u)) - R_v(u)$

Entretanto, é lógico que a projeção de um vetor e a projeção de sua reflexão, na mesma base, são iguais. Portanto:

$R_v(R_v(u)) = 2P_v(u) - R_v(u)$

Que é exatamente a relação que obtemos quando isolamos o $u$ na primeira equação, logo:

$R_v(R_v(u)) = u$

Use esta ideia para construir testes (aleatórios) para a sua função de reflexão.

<!-- grade_id: cell-bcce0db9283d8138, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE

for i in range(1, 50):
    v = np.random.rand(i*5)
    u = np.random.rand(i*5)

    r1 = refl(v)(u)
    r2 = refl(v)(r1)

    assert np.allclose(u, r2)
```

# Parte 2: Integrais interpolatórias

Vimos, na lista 3, que dados os pontos $x_i$ no intervalo $[-1,1]$,
podemos criar uma "regra de integração" para uma função $f$ que é equivalente a:
1. Calcular o polinômio interpolador de $f$ nos nós $x_i$
2. Calcular a integral deste polinômio no intervalo $[-1,1]$.

## Questão 7: Pontos de Chebyshev

Há duas famílias de pontos de Chebyshev:
1. Interiores: da forma $\cos\left( \frac{2k+1}{2n} \pi \right)$, para $k$ de $0$ a $n$ (exclusive)

2. De bordo: da forma $\cos\left( \frac{k}{n} \pi \right)$, para $k$ de $0$ a $n$ (inclusive)

<!-- grade_id: cell-df2b92ea394723d6, points: 0, grade: false, solution: true -->
```python
def cheby1(n):
    # YOUR CODE HERE
    raise NotImplementedError()

def cheby2(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-efd330570d34aef6, points: 1, grade: true, solution: false -->
```python
assert len(cheby1(5)) == 5
assert len(cheby2(5)) == 6
```

<!-- grade_id: cell-3ecf919acbecc8b3, points: 1, grade: true, solution: false -->
```python
c5 = cheby1(15)
assert np.all(c5[1:] < c5[:-1])

c5 = cheby2(25)
assert np.all(c5[1:] < c5[:-1])
```

Faça um desenho destes pontos, e explique a nomenclatura.

<!-- grade_id: cell-ff43b16e64da3292, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

ax = plt.gca()
```

<!-- grade_id: cell-b3f2fabcda1efd33, points: 1, grade: true, solution: false -->
```python
assert len(ax.legend().texts) == 2
assert len(ax.get_title()) >= 10
```

Explique aqui

<!-- grade_id: cell-e9de80ad3b063813, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

```python
ax = None
```

## Questão 8: Integrando polinômios

Escreva a função que, para um polinômio $P$, dado por seus coeficientes em ordem crescente,
calcula a integral de $P$ no intervalo $[-1,1]$.

<!-- grade_id: cell-007fb0309edaf766, points: 0, grade: false, solution: true -->
```python
def base_int(p):
    # YOUR CODE HERE
    valor = 0
    for grau, coef in enumerate(p):
        valor_cima = coef/(grau+1)
        valor_baixo = (coef/(grau+1))*(-1)**(grau+1)
        valor += valor_cima - valor_baixo
    return valor
base_int([1, 2, 3])
```

Output:
```
4.0
```

<!-- grade_id: cell-582995aaa41e817f, points: 1, grade: true, solution: false -->
```python
assert base_int([1,2,3]) == 4
```

<!-- grade_id: cell-f074ce24df68da95, points: 1, grade: true, solution: false -->
```python
assert np.isclose(base_int([5,4,3,2,1]), 12.4)
```

## Questão 9: Integração interpolatória

Escreva uma função que recebe um conjunto $n$ de pontos $x_i$ do intervalo $[-1,1]$,
e calcula pesos $w_i$ tais que

$$ \int_{-1}^1 P(x) \, dx = \sum_{i=1}^n w_i P(x_i) $$

para todos os polinômios $P$ de grau menor do que $n$.

<!-- grade_id: cell-fe56e1c6bb2076a1, points: 0, grade: false, solution: true -->
```python
def pesos(xs):
    xs = np.asarray(xs)
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-60a4c4249638d7b9, points: 1, grade: true, solution: false -->
```python
assert np.allclose(pesos([-1, 1]), [1, 1])
```

<!-- grade_id: cell-547ea592a0905047, points: 1, grade: true, solution: false -->
```python
assert np.allclose(pesos([-1, 0, 1]), [1/3, 4/3, 1/3])
```

O que estes dois primeiros testes verificam?

<!-- grade_id: cell-e34d20d7d99b7a01, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: cell-77c8235dd3c11971, points: 1, grade: true, solution: false -->
```python
w4 = pesos(cheby1(4))
assert np.allclose(w4, [0.26429774, 0.73570226, 0.73570226, 0.26429774])
```

## Questão 10: Integrando funções

Vamos usar os mesmos pesos para integrar funções que não são polinômios.

<!-- grade_id: cell-cf6cee2e0eeb41e9, points: 1, grade: true, solution: false -->
```python
xs = cheby1(4)
w4 = pesos(xs)

assert 2e-4 <= np.dot(w4, np.cos(xs)) - (np.sin(1) - np.sin(-1)) <= 2e-3
assert 2e-4 <= np.dot(w4, np.exp(xs)) - (np.exp(1) - np.exp(-1)) <= 2e-3
```

O que estes testes mostram?

<!-- grade_id: cell-0ca9baada089e7e4, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: cell-a73bf1cdc1eec899, points: 1, grade: true, solution: false -->
```python
xs = cheby2(10)
w10 = pesos(xs)

assert np.abs(np.dot(w10, np.cos(xs)) - (np.sin(1) - np.sin(-1))) < 1e-12
assert np.abs(np.dot(w10, np.exp(xs)) - (np.exp(1) - np.exp(-1))) < 1e-12
```

O que os estes testes mostram, em comparação com os anteriores?
Isso é esperado?

<!-- grade_id: cell-a69a1092b0836c7d, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## Questão 11: Evolução da precisão

Faça um gráfico do erro de integração da função $e^{2x}$ no intervalo $[-1, 1]$
em função do número de pontos de Chebyshev utilizados.

Deve haver duas curvas: uma para os pontos de tipo 1, outra para os pontos de tipo 2.

<!-- grade_id: cell-440689db152d8fc1, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

ax = plt.gca()
```

<!-- grade_id: cell-afceb523f517ff3b, points: 1, grade: true, solution: false -->
```python
assert len(ax.lines) == 2
assert len(ax.legend().texts) == 2
```

<!-- grade_id: cell-c18866a7b661cde3, points: 1, grade: true, solution: false -->
```python
assert len(ax.get_title()) >= 10
assert len(ax.get_xlabel()) >= 5
```

O que você percebe quanto à velocidade com que os erros tendem a zero?

<!-- grade_id: cell-aa3ec382472f6555, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

## Questão 12: Integradores

Escreva a função `make_integrator` que recebe `xs`, os pontos $x_i$ no intervalo $[-1,1]$,
e retorna uma função `int_interp(f, a, b, n)` que
- Recebe uma função $f$, um intervalo $[a,b]$ e um número de subdivisões $n$
- Divide o intervalo em $n$ subintervalos de mesmo tamanho
- Calcula a integral de $f$ em cada um destes subintervalos usando pontos correspondentes (por translação e homotetia) aos pontos $x_i$
- E retorna a soma como uma aproximação da integral de $f$ no intervalo $[a,b]$.

<!-- grade_id: cell-f5bc035273e90cc1, points: 0, grade: false, solution: true -->
```python
def make_integrator(xs):
    xs = np.asarray(xs, dtype=float)
    ws = pesos(xs)
    def int_interp(f, a, b, n):
        # YOUR CODE HERE
        raise NotImplementedError()

    return int_interp
```

<!-- grade_id: cell-7758c534233a04c7, points: 1, grade: true, solution: false -->
```python
f2 = make_integrator([1/2, -1/2])
assert 1e-4 < np.abs(f2(np.exp, 2, 5, 200) - (np.exp(5) - np.exp(2))) < 1e-3
```

<!-- grade_id: cell-4cb4ab992ed1a1a2, points: 1, grade: true, solution: false -->
```python
f5 = make_integrator([0,1/2,-1/2,1,-1])
assert 5e-14 < np.abs(f5(np.exp, 2, 5, 80) - (np.exp(5) - np.exp(2))) < 5e-13
```

Comente os dois testes acima, e explique porque eles são uma comparação "justa" dos métodos.

<!-- grade_id: cell-8bb198e113d3e02d, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

Calcule os erros de integração da função $e^{2x}$ no intervalo $[-1,4]$,
para o método que usa 10 pontos de Chebyshev,
em função do número de subintervalos.

<!-- grade_id: cell-9b46129551024d94, points: 1, grade: true, solution: true -->
```python
fcheby = make_integrator(cheby1(10))
# YOUR CODE HERE
raise NotImplementedError()

ax = plt.gca()
```

<!-- grade_id: cell-1dbd59cc850a3a80, points: 1, grade: true, solution: false -->
```python
assert len(ax.lines) == 1

assert len(ax.get_title()) >= 10
assert len(ax.get_xlabel()) >= 5
```

Qual a velocidade de convergência deste método?

<!-- grade_id: cell-ceb1836b149d5372, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE